// variables6.rs
//
// Execute `rustlings hint variables6` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE
//Correction
const NUMBER : i32 = 3; // j'ai  ajouté le type i32 à la constante NUMBER pour indiquer que c'est un entier 
fn main() {
    println!("Number {}", NUMBER);
}
//explication
//Le code  ne compile pas car la constante NUMBER est définie comme une valeur entière, mais elle n'a pas de type explicite. 
