// variables2.rs
//
// Execute `rustlings hint variables2` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

//Correction
fn main() {
    let x = 10 ; // x a été initialisé avec une valeur.
    if x == 10 {
        println!("x is ten!");
    } else {
        println!("x is not ten!");
    }
}
//explication
//  La variable x est déclarée, mais elle n'est pas initialisée avant d'être utilisée dans la condition if