// variables1.rs
//
// Make me compile!
//
// Execute `rustlings hint variables1` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE
//Correction
fn main() {
    let x: i32 = 5;
    println!("x has the value {}", x);
}

// explication 
// la variable x n'avait pas été déclarée j'ai donc déclaré avec let  la variable ainsi que le type