// variables3.rs
//
// Execute `rustlings hint variables3` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE
//Correction
fn main() {
    let x: i32 = 6 ; //x a été initialisé avec une valeur 6
    println!("Number {}", x);
}
//explication
//  la variable x est déclarée, mais elle n'est pas initialisée avant d'être utilisée dans println . j'ai du affecté une valeur a x 
