// functions1.rs
//
// Execute `rustlings hint functions1` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE
//Correction
fn call_me() {
    println!("hello rust");
}
fn main() {
    call_me();
}

//explication
// le code ne compile pas car la fonction call_me n'est définie nulle part dans le code mais est appelée dans la fonction main
// j'ai juste definie la fonction call me 