// primitive_types1.rs
//
// Fill in the rest of the line that has code missing! No hints, there's no
// tricks, just get used to typing these :)

// I AM NOT DONE

//correction
fn main() {
    // Booleans (`bool`)

    let is_morning = true;
    if is_morning {
        println!("Good morning!");
    }

    let is_evening = true; // Finish the rest of this line like the example! Or make it be false!
    if is_evening {
        println!("Good evening!");
    }
}

//explication
//La deuxième partie est incomplète car il manquue le type de la variable qui est boolen.
// pour la completer j'ai definie une variable is_evening de type bool et de lui attribuer une valeur true.l'instruction if vérifie si is_evening est vrai, et si c'est le cas, elle affiche le message "Good evening!".
